<div style="background:#222222; color:#ffffff; padding:20px">
<h1 align="center">Guided ML With The Iris Dataset</h1>

<h2 align="center" tyle="color:#01ff84" >Learning type | Activity type | Objective |</h2>
<h2 align="center">| Supervised | Multiclass classification | Identify a flower's class |</h2>


<div>

Contents:

1. Loading the data
2. Setting up supervised learning problem (selecting features)
3. Creating a first model
    - Creating train and test datasets
    - Normalizing train and test
    - Fitting and predicting
4. Evaluate the frist model predictions
5. Crossvalidation of the model
6. Creating an end to end ML pipeline
    - Train/Test Split
    - Normalize
    - Crossvalidations
    - Model
    - fitting and predicting

## Instructions with NBGrader removed

Complete the cells beginning with `# YOUR CODE HERE` and run the subsequent cells to check your code.


Contents:
1. Loading the data
2. Setting up supervised learning problem (selecting features)
3. Creating a first model
    - Creating train and test datasets
    - Normalizing train and test
    - Fitting and predicting
4. Evaluate the frist model predictions
5. Crossvalidation of the model
6. Creating an end to end ML pipeline
    - Train/Test Split
    - Normalize
    - Crossvalidations
    - Model
    - fitting and predicting

## About the dataset

[Iris](https://archive.ics.uci.edu/ml/datasets/iris) is a well-known multiclass dataset. It contains 3 classes of flowers with 50 examples each. There are a total of 4 features for each flower.

![](./classic-datasets/images/Iris-versicolor-21_1.jpg)

## Package setups

1. Run the following two cells to initalize the required libraries. 

```python
# to debug package errors
import sys
sys.path
sys.executable

```

Output:
```
'C:\\Users\\ritth\\software\\anaconda\\envs\\strive\\python.exe'
```

```python
# Import needed packages
# You may add or remove packages should you need them
import numpy as np
import matplotlib
import matplotlib.pyplot as plt
import seaborn as sns
import pandas as pd
from sklearn import datasets
from sklearn import model_selection
from sklearn.preprocessing import StandardScaler
from sklearn import preprocessing
from sklearn.linear_model import LogisticRegression
from sklearn.model_selection import train_test_split, cross_validate
from sklearn.pipeline import make_pipeline

# Display plots inline and change plot resolution to retina
%matplotlib inline
%config InlineBackend.figure_format = 'retina'
# Set Seaborn aesthetic parameters to defaults
sns.set()
```

## Step 1: Loading the data

1. Load the iris dataset using ```datasets.load_iris()```
2. Investigate the data structure with ```.keys()```
3. Construct a dataframe from the dataset
4. Create a 'target' and a 'class' column that contains the target names and values
5. Display a random sample of the dataframe 

<!-- grade_id: cell-c92cca7755a29a2c, points: 0, grade: false, solution: true -->
```python
def load_data():
    #load the dataset
    #return the dataset
    
    return datasets.load_iris()
    
```

<!-- grade_id: cell-6955aa3f0eb9b484, points: 1, grade: true, solution: false -->
```python
assert load_data()['data'].shape == (150,4)
```

```python
# to check the dicitionary keys
load_data().keys()
```

Output:
```
dict_keys(['data', 'target', 'frame', 'target_names', 'DESCR', 'feature_names', 'filename', 'data_module'])
```

```python
dataset = load_data()
load_data().feature_names
```

Output:
```
['sepal length (cm)',
 'sepal width (cm)',
 'petal length (cm)',
 'petal width (cm)']
```

```python
dataset['target_names']
```

Output:
```
array(['setosa', 'versicolor', 'virginica'], dtype='<U10')
```

<!-- grade_id: cell-054f7e75dfc37e15, points: 0, grade: false, solution: true -->
```python
def dataset_to_pandas():
    #put the dataset into a pandas DF using the feature names as columnsç
    #rename the column name so the dont include the '(cm)'
    #add 2 columns one with the target and another with the target_names
    
    # create data frame
    df = pd.DataFrame(dataset['data'], columns = dataset['feature_names'])

    # clean column (cm)
    df['columns'] = df.columns.str.replace('\(cm\)', "").str.strip()

    # # getting target
    # df['target'] = dataset['target']

    # # setting the values
    # df['class'] = df['target'].map({0:'setosa', 1:'versicolor', 2:'virginica'})
    return 

dataset_to_pandas()
```

Output:
```
File ~\software\anaconda\envs\strive\lib\site-packages\pandas\core\frame.py:3764, in DataFrame._sanitize_column(self, key, value, broadcast)
   <a href='file:///c%3A/Users/ritth/software/anaconda/envs/strive/lib/site-packages/pandas/core/frame.py?line=3758'>3759</a>     value = sanitize_index(value, self.index)
   <a href='file:///c%3A/Users/ritth/software/anaconda/envs/strive/lib/site-packages/pandas/core/frame.py?line=3760'>3761</a> elif isinstance(value, Index) or is_sequence(value):
   <a href='file:///c%3A/Users/ritth/software/anaconda/envs/strive/lib/site-packages/pandas/core/frame.py?line=3761'>3762</a> 
   <a href='file:///c%3A/Users/ritth/software/anaconda/envs/strive/lib/site-packages/pandas/core/frame.py?line=3762'>3763</a>     # turn me into an ndarray
-> <a href='file:///c%3A/Users/ritth/software/anaconda/envs/strive/lib/site-packages/pandas/core/frame.py?line=3763'>3764</a>     value = sanitize_index(value, self.index)
   <a href='file:///c%3A/Users/ritth/software/anacond
```

<!-- grade_id: cell-b18a01895567bcda, points: 2, grade: true, solution: false -->
```python
df = dataset_to_pandas()
assert df['sepal length'].shape == (150,)
assert df['sepal width'].shape == (150,)
assert df['petal length'].shape == (150,)
assert df['petal width'].shape == (150,)
assert df['target'].shape == (150,)
assert df['class'].shape == (150,)

```

Output:
```
TypeError                                 Traceback (most recent call last)
c:\Users\ritth\code\Strive\Strive-Exercises\Chapter 02\01. Intro to ML\Part 1\exercises\ml1.ipynb Cell 16' in <cell line: 2>()
      <a href='vscode-notebook-cell:/c%3A/Users/ritth/code/Strive/Strive-Exercises/Chapter%2002/01.%20Intro%20to%20ML/Part%201/exercises/ml1.ipynb#ch0000014?line=0'>1</a> df = dataset_to_pandas()
----> <a href='vscode-notebook-cell:/c%3A/Users/ritth/code/Strive/Strive-Exercises/Chapter%2002/01.%20Intro%20to%20ML/Part%201/exercises/ml1.ipynb#ch0000014?line=1'>2</a> assert df['sepal length'].shape == (150,)
      <a href='vscode-notebook-cell:/c%3A/Users/ritth/code/Strive/Strive-Exercises/Chapter%2002/01.%20Intro%20to%20ML/Part%201/exercises/ml1.ipynb#ch0000014?line=2'>3</a> assert df['sepal width'].shape == (150,)
      <a href='vscode-notebook-cell:/c%3A/Users/ritth/code/Strive/Strive-Exercises/Chapter%2002/01.%20Intro%20to%20ML/Part%201/exercises/ml1.ipynb#ch0000014?line=3'>4</a> asser
```

### Question
Find the X and y values we're looking for. Notice that y is categorical and thus, we could **one-hot encode it** if we are looking at **class** or we can just pick **target**. In order to one hot encode we have  to re-shape `y` it using the **.get_dummies** function. 

## For the purpose of this exercise, do not use hot encoding, go only for target but think about if you have to drop it somewhere or not...

```python
pd.get_dummies
```

Output:
```
<function pandas.core.reshape.reshape.get_dummies(data, prefix=None, prefix_sep='_', dummy_na=False, columns=None, sparse=False, drop_first=False, dtype=None) -> 'DataFrame'>
```

```python
df_iris = dataset_to_pandas()

```

Output:
```
NameError                                 Traceback (most recent call last)
c:\Users\ritth\code\Strive\Strive-Exercises\Chapter 02\01. Intro to ML\Part 1\exercises\ml1.ipynb Cell 19' in <cell line: 1>()
----> <a href='vscode-notebook-cell:/c%3A/Users/ritth/code/Strive/Strive-Exercises/Chapter%2002/01.%20Intro%20to%20ML/Part%201/exercises/ml1.ipynb#ch0000016?line=0'>1</a> df_iris = dataset_to_pandas()

NameError: name 'dataset_to_pandas' is not defined
```

```python
# check type
df_iris.dtype()
```

```python
# see class
df_iris['class']
```

```python
pd.get_dummies(df_iris)
```

<!-- grade_id: cell-fd162bfaf4394e79, points: 0, grade: false, solution: true -->
```python
# virginica close to versicolor BAD
#0 - setosa
#1 - versicolor
#2 - virginica 

# Same distance between all the classes GOOD
#[1,0,0] - setosa
#[0,1,0] - versicolor
#[0,0,1] - virginica 


def ohe():
    return pd.get_dummies(df_iris)
```

<!-- grade_id: cell-423f24383553b71a, points: 1, grade: true, solution: false -->
```python
ohe_data = ohe()

assert ohe_data.shape == (150,8)
```

## Step 2: Setting up supervised learning problem (selecting features)

Feature selection is an essential step in improving a model's perfromance. In the first version of the model we will use the **'sepal length'** and **'sepal width'** as predicting features. Later we will see the effect of adding additional features.

1. Assign the values of the 'target' to Y as a numpy array
2. Assign the remaining feature values to X as a numpy array
3. Check the shape of X and Y. Check the first few values.
    - Can we confirm our X and Y are created correctly?

```python
# use reshape to get 1
# run the code by step by step
np.array(df_iris.target).reshape(-1,1).shape
```

```python
# df_iris[['sepal length', 'sepal width']].values
# values and to_numpy()

df_iris[['sepal length', 'sepal width']].to_numpy().shape
```

<!-- grade_id: cell-c0cf28fd18447a0b, points: 0, grade: false, solution: true -->
```python
def target_to_numpy():
    return np.array(df_iris.target)

def data_to_numpy():
     return df_iris[['sepal length', 'sepal width']].to_numpy()
    
```

<!-- grade_id: cell-aaf03dfa38041beb, points: 1, grade: true, solution: false -->
```python
Y = target_to_numpy()
X = data_to_numpy()
assert isinstance(Y, np.ndarray)
assert isinstance(X, np.ndarray)
assert X.shape == (150,2)
```

```python
#your code here
X = df_iris[['sepal length', 'sepal width']].values
print(X.shape)
X[:5]
```

## Step 3: Creating the first model

In lecture we learned about creating a train and test datasets, normalizing, and fitting a model. In this step we will see how to build a simple version of this.

We have to be careful when constructing our train and test datasets. First, when we create train and test datasets we have to be careful that we always have the same datapoints in each set. Otherwise our results won't be reproduceable or we might introduce a bias into our model.

We also need to be attentive to when we normalize the data. What would be the effect of normalizing the data (i.e. with StandardScaler to a range between 0 - 1) before we create our train and test sets? Effectively we would use information in the test set to structure the values in the training set and vice versa. Therefore normalizing train and test independently is the preferred method.

1. Create X_train, X_test, Y_train, Y_test using ```train_test_split()``` with an 80/20 train/test split. Look in the SKLearn documentation to understand how the function works.
    - Inspect the first few rows of X_train.
    - Run the cell a few times. Do the first few rows change?
    - What option can we use in ```train_test_split()``` to stop this from happening?
2. Normalize the train and test datasets with ```StandardScaler```
    - We can fit the transform with ```.fit()``` and ```.transform()``` to apply it. Look in the documentation for an esample of how to do this.
    - Does it make sense to normalize Y_train and Y_test?
3. Initalize a ```LogisticRegression()``` model and use the ```.fit()``` method to initalize the first model.
    - We will pass the X_train and Y_train variables to the ```.fit()``` method.
    - Once the model is fit, use the ```.predict()``` with the X_test and save the output as predictions.

<!-- grade_id: cell-586cc3c4a3d849fa, points: 0, grade: false, solution: true -->
```python
#split train and test data 80/20 use random_state=0, this is the state used for when the asserts were written
#your code here
X_train, X_test, Y_train, Y_test = 0,0,0,0

# test with train_test_split
X_train, X_test, Y_train, Y_test = train_test_split(x, y, test_size= 0.2, random_state=0)

print(X_train.shape)
print(Y_train.shape)
print(X_test.shape)
print(Y_test.shape)
```

<!-- grade_id: cell-4415454a211f5895, points: 1, grade: true, solution: false -->
```python
assert X_train.shape == (120,2)
assert Y_train.shape == (120,)
assert X_test.shape  == (30,2)
assert Y_test.shape  == (30,)
```

```python
second_scalar = StandardScaler()

# fit help to check mean and standard deviation
second_scalar.fit(X_train)

# transform help to convert fit into normalisation
second_scalar.transform(X_train)
```

<!-- grade_id: cell-d94703390548f2df, points: 0, grade: false, solution: true -->
```python
# normalize the dataset
# create and fit the scaler object on the training data

scalar = StandardScaler()
# in simple way we can write as
scalar.fit_transform(X_train)

X_train[:5]
```

```python
# to check test we use only transform
scalar.transform(X_test)

X_test[:5]
```

<!-- grade_id: cell-fd525045d90798a3, points: 2, grade: true, solution: false -->
```python
assert np.amin(X_train) >= -2.5
assert np.amax(X_train) <= 3.2
assert np.amin(X_test) >= -2
assert np.amin(X_test) <= 2.75
```

<!-- grade_id: cell-a91abe39c7261df3, points: 0, grade: false, solution: true -->
```python
# initalize and fit with Logistic Regression
prediction = 0

# initalize the logistic regressor
# make predictions
model = LogisticRegression()



```

<!-- grade_id: cell-61e5bc16122b7220, points: 2, grade: true, solution: false -->
```python
assert predictions.shape == (30,)
```

## Step 4: Evaluate the frist model's predictions

We will learn more about how to evaluate the performance of a classifier in later lessons. For now we will use % accuracy as our metric. It is important to know that this metric only helps us understand the specific performance of our model and not, for example, where we can improve it, or where it already perfoms well.

1. Use ```.score()``` to evaluate the performance of our first model.

<!-- grade_id: cell-a19f9db86d767c57, points: 0, grade: false, solution: true -->
```python
score = 0
# evaluating the performace of our first model

```

<!-- grade_id: cell-300a4e426c8fcfa1, points: 1, grade: true, solution: false -->
```python
assert score >=0.7
```

<!-- grade_id: cell-f7dbd3aa1419397f, points: 1, grade: true, solution: false -->
```python
assert score >=0.72
```

<!-- grade_id: cell-3b6954a5b2f244cd, points: 1, grade: true, solution: false -->
```python
assert score >=0.73
```

## Step 5: Crossvalidation of the model
Our first model achived ~90% accruacy. This is quite good. How do we know it is reproducable? If we run the model again and our performance is 85% which is correct? And what about improving our model? Can you think of one thing we can do to potentially improve the model?

#### Crossvalidation
Corssvalidation is when we create multiple X and Y datasets. On each dataset we train and fit the model. We then average the results and return a 'crossvalidated' accruacy.

1. Initalize a new version of the model you trained above with the same paramters.
2. Use ```cross_validate()``` to run the model with 5 crossvalidation folds. 

<!-- grade_id: cell-49d1242d4949806f, points: 0, grade: false, solution: true -->
```python
# model with cross validation

# cross validate the training set
clf_cv = 0
CV = 0

def print_scores(cv):
    # print out cross validation scores
    [print('Crossvalidation fold: {}  Accruacy: {}'.format(n, score)) for n, score in enumerate(cv['test_score'])]
    # print out the mean of the cross validation
    print('Mean train cross validation score {}'.format(cv['test_score'].mean()))
    
print_scores(cv)
```

<!-- grade_id: cell-cbdb686462858d91, points: 4, grade: true, solution: false -->
```python
assert len(cv['test_score']) == 5
assert max(cv['test_score']) >= 0.85
assert min(cv['test_score']) >= 0.69
assert cv['test_score'].mean() >= 0.77
```

## Step 6: Creating an end to end ML pipeline
Congraulations you've trained, crossvalidated, predicted, and evaluated your frist classifier. Now that you understand the basic steps we will look at a way to combine all these steps together.

Before we go further think about what you would have to do if you wanted to change the model. Intalize a new model, change the vairables, redo the cross validation...etc. Seems like a lot. And when we have to change lots of code it is easy to make mistakes. And what if you wanted to try many models and see which one performed best? Or try changing many different features? How could you do it without writing each one out as we have?

The solution is to use SKLearn's pipeline class. A pipeline is an object that will execute the various steps in the machine learning process. We can choose what elements we want in the pipeline and those that we do not. Once setup, we can rapidly change models, or input data and have it return our results in an ordered way.


1. Initalize a scaler and a classifer object like we did previously.
2. Use the ```make_pipeline()``` function to construct a transofmraiton pipeline for the scaler and the classifier
3. Input the pipeline object to the cross_validator and evaluate with 5 folds.
4. Print out your results (hint: make a function for repetitve tasks like printing)

<!-- grade_id: cell-82685348a4b1d16a, points: 0, grade: false, solution: true -->
```python
# define the scaler
# define the classifier
# make the pipeline
# run the cross validation
# print results
scaler = 0
classifier = 0
pipe = 0
scores = 0

```

<!-- grade_id: cell-0f5c78a78336a3c0, points: 4, grade: true, solution: false -->
```python
assert type(pipe) == type(make_pipeline(scaler, classifier))
assert len(cv['test_score']) == 5
assert max(cv['test_score']) >= 0.83
assert min(cv['test_score']) >= 0.69
assert cv['test_score'].mean() >= 0.74
```

## Challenge Exercise

In this notebook we only used two features to predict the class of the flower. We also did not do any hypter parameter tuning. The challenge is to impove the prediction results. Some ideas we can try:
1. Add features to the input and run the cross validation pipeline
2. Investigate how to use ```GridSearchCV```, a powerful funtion that searches through hyperparmetrs and does cross validation.
    - Hint: Input the pipeline directly into GridSearchCV
3. Try a different models like RandomForest or SVM.

<!-- grade_id: cell-f1ba9880447e80ec, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-71c06d508e0294b0, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-230a4548c4b43eb7, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-be4475eea0463f76, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```

<!-- grade_id: cell-b31a03b36b3beb45, points: 0, grade: false, solution: true -->
```python
# YOUR CODE HERE
raise NotImplementedError()
```